```python
# Initialize Otter
import otter
grader = otter.Notebook("lab15.ipynb")
```

# Lab 15: Data Ethics Concepts

<!-- BEGIN QUESTION -->

**What is your name?**

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 1. Laws and Regulations

### 1.1 The European General Data Protection Regulation (GDPR)

<!-- BEGIN QUESTION -->

**Question 1.1.1.**

Under GDPR, how long can a business save personally identifying information in an archive?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.1.2.**

Under GDPR, can a business collect and process personal data in order to save someone's life?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.1.3.**

What is the right to erasure?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 1.2 The California Privacy Rights Act (CPRA) of 2020

<!-- BEGIN QUESTION -->

**Question 1.2.1.**

Under the California Privacy Rights Act, how long can a business save personally identifying information in an archive?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.2.2.**

Under the California Privacy Rights Act, do people have the right to stop businesses from using their race?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 1.3 Privacy Rights in Ohio

Neither the GDPR nor the CPRA apply in Ohio.

<!-- BEGIN QUESTION -->

**Question 1.3.1.**

What privacy right do people in Ohio have?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 1.4 Section 230 of the Communications Decency Act (CDA)

<!-- BEGIN QUESTION -->

**Question 1.4.1.**

Under section 230 of the CDA, can you host a website on which other people post lies, post threats, or try to defraud others without you getting in trouble?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.4.2.**

What are the benefits of having section 230 of the CDA as the law?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Unfair Discrimination and Social Bias

<!-- BEGIN QUESTION -->

**Question 2.1.**

What is *Algorithmic Bias*?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.2.**

What is *redlining*?
Did redlining happen in the area where you grew up?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3. Transparency of Methods and Analysis

<!-- BEGIN QUESTION -->

**Question 3.1.**

Suppose the data you collect contradicts the central claim in your doctoral dissertation.
If you are not an ethical person, what could you do so you still get to graduate?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.2.**

Suppose the data you collect contradicts the central claim in your doctoral dissertation.
If you are an ethical person, what could you do so you still get to graduate?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Sampling Bias

<!-- BEGIN QUESTION -->

**Question 4.1.**

Suppose that student evaluations of a course only include students who finished the course.
Which type of sampling bias would this introduce?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Submission

<!-- BEGIN QUESTION -->

**Acknowledgements:**
If you received help from any other person, then describe here what help you received and from whom. 
If not, then say 'None'.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

It is forbidden to: 
* Use the work or programs from students who took this class (in any of its versions at any university).
* Use websites that claim to have solutions for this class (in any of its versions at any university).
* Directly copy from a classmate.

**Did you do anything forbidden while working on this assignment?**

Assign `did_forbidden` to either `True` or `False`.

```python
do_forbidden = ...
```

```python
grader.check("q_forbidden")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

After saving and exporting, submit the .zip file in BlackBoard

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q_forbidden (points: None)
```
>>> do_forbidden in [True,False]
True
```
