```python
# Initialize Otter
import otter
grader = otter.Notebook("hw15.ipynb")
```

# Homework 15: Data Ethics Scenarios

<!-- BEGIN QUESTION -->

**What is your name?**

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 1. Home-made Scenarios

### 1.1 Hiring

A company wants to hire people that it thinks will stay and be successful at that company. 
It collects data from the applications of people it hired 5 years ago, and labels each as successful, unsuccessful, or left.
It then constructs a kNN classifier using these three classes.

<!-- BEGIN QUESTION -->

**Question 1.1.1.**

What benefit could the company expect from using this classifier when making hiring decisions?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.1.2.**

How could using this classifier when making hiring decisions reinforce social biases?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.1.3.**

What could the company do to gain the benfits of using a classifier without reinforcing social biases?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 1.2 MATH Placement

When students enter a university, they are given a Mathematics placement level, which determines which MATH courses they can enroll in initially. For example it determines whether they can start in Calculus or need to first take Pre-Calculus.

For this simplified example, assume that the only two classes are Pre-Calculus and Calculus.

<!-- BEGIN QUESTION -->

**Question 1.2.1**

University A requires students to take the [ACT](https://www.act.org/) standardized test and places them into Calculus if they score at least 30.

What are the advantages and disadvantages of this method?
Could it be biased toward or against some group?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.2.2.**

University B places student based only on their High School GPA.
If they have at least 3.5, then they are placed into Calculus.

What are the advantages and disadvantages of this method?
Could it be biased toward or against some group?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.2.3.**

University C requires students to take the [ACT](https://www.act.org/) standardized test but also uses their High School GPA when placing them.

Using recent students, they gathered data on their ACT score, GPA, and grade in Calculus. It looked something like

| ACT   | GPA    | Grade |
| ------ | ------ | ----- |
| 24     | 3.6    | B+    |
| 32     | 2.1    | D     |
| ... | ... | ... |

They constructed a multiple-regression model that used both `ACT` and `GPA` to predict the Grade (converted to 0-4 scale).
(Thus, for example, a higher GPA could compensate for a lower ACT score.)
They then place students into Calculus if their predicted grade is at least C (2.00).

What are the advantages and disadvantages of this method?
Could it be biased toward or against some group?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.2.4.**

University D requires students to take the [ACT](https://www.act.org/) standardized test but also considers other information in their application (GPA, zip code, rating of their High School, gender, ...).

Using recent students, they gathered data on their ACT score, other application information, and grade in Calculus. 
They constructed a multiple-regression model that used all the available information to predict the Grade (converted to 0-4 scale).
They then place students into Calculus if their predicted grade is at least C (2.00).

What are the advantages and disadvantages of this method?
Could it be biased toward or against some group?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.2.5.**

If University D removed gender from their data, what other variables could possibly be used as a proxy for gender? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Borrowed Scerarios

Now we'll consider some [exercises](https://mdsr-book.github.io/mdsr2e/ch-ethics.html#ethics-exercises) from [Modern Data Science with R](https://mdsr-book.github.io/mdsr2e/).
To avoid copyright violations, we will just give the problem number and start.

<!-- BEGIN QUESTION -->

**Question 2.1.**

[Problem](https://mdsr-book.github.io/mdsr2e/ch-ethics.html#ethics-exercises) 2 (Medium): A data scientist compiled data ... 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.2.**

[Problem](https://mdsr-book.github.io/mdsr2e/ch-ethics.html#ethics-exercises) 5 (Medium): A reporter carried out a clinical trial of chocolate ...

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Submission

<!-- BEGIN QUESTION -->

**Acknowledgements:**
If you received help from any other person, then describe here what help you received and from whom. 
If not, then say 'None'.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

It is forbidden to: 
* Use the work or programs from students who took this class (in any of its versions at any university).
* Use websites that claim to have solutions for this class (in any of its versions at any university).
* Directly copy from a classmate.

**Did you do anything forbidden while working on this assignment?**

Assign `did_forbidden` to either `True` or `False`.

```python
do_forbidden = ...
```

```python
grader.check("q_forbidden")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

After saving and exporting, submit the .zip file in BlackBoard

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q_forbidden (points: None)
```
>>> do_forbidden in [True,False]
True
```
